# Part 1: Aliasing Theory Quiz

This is a quick quiz to test your understanding of Sampling and Aliasing fundamentals. Each question is to be answered with a numerical value (`float` or `int`), and you can check your answer by running the test cell directly below the cell where you write your answer.

Relevant reading materials:
* Chapter 2 in Understanding Digital Signal Processing
* Chapter 12.2 & 12.3.1-2 in Understanding Digital Signal Processing
* Lecture notes on Sampling and Aliasing

## a) 
A signal $x_1(t)=\sin⁡(2\pi \cdot f \cdot t)$ with frequency $f = 40 \text{ Hz}$ is sampled with a sampling frequency $f_s = 200 \text{ Hz}$. What is the normalized angular frequency $\hat{\omega}$ for the sampled signal $x_1[n]$?

<!-- grade_id: cell-7866267f359241f7, points: 0, grade: false, solution: true -->
```python
import numpy as np
f = 40 # Hz
fs = 200 #Hz
w_norm = 2*np.pi * f/fs  # Replace with correct value
print(f"{w_norm = }")
```

Output:
```
w_norm = 1.2566370614359172
```

<!-- grade_id: cell-457f0b8a3ada5a61, points: 4, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(w_norm), 3)).encode('utf-8')+b'e585b').hexdigest() == '99dde87491837c22915c02f47bc4cd71367274da', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## b)
A digital signal $x_2[n] =\sin(\hat{\omega}_2 \cdot n)$ with angular frequency $\hat{\omega}_2 = \frac{\pi}{5}$ is converted to an analog signal $x_2(t)$ usiong a sampling frequency $f_s = 200 \text{ Hz}$, where we assume ideal reconstruction. What is the apparent frequency $f_2$ of the reconstructed signal $x_2(t)$?

<!-- grade_id: cell-2f03a8f12caebb02, points: 0, grade: false, solution: true -->
```python
w_norm = np.pi / 5
fs = 200

f_2 = w_norm*fs/(2*np.pi)
print(f"{f_2 = }")# Replace with correct value
```

Output:
```
f_2 = 20.0
```

<!-- grade_id: cell-eebc2ed015c7ddeb, points: 4, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(f_2), 0)).encode('utf-8')+b'7031e').hexdigest() == 'd5e0590fd3b888c1f826570f90ce72fc7e085669', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## c) 
What is the maximum frequency which can be represented with a digital signal when the sampling frequency is $f_s = 200Hz$?

<!-- grade_id: cell-ad8f77a7dbedd360, points: 0, grade: false, solution: true -->
```python
fs = 200 # Hz
f_max = fs/2  # Replace with correct value
print(f"{f_max = }")
```

Output:
```
f_max = 100.0
```

<!-- grade_id: cell-fca70d1e16a617c2, points: 4, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(f_max), 0)).encode('utf-8')+b'78f94').hexdigest() == '210058739458b370a60612ec245f8a1ec460a7f6', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## d)
What is the absolute minimum sampling frequency of a signal which permits successfull representation of a sinusoid with frequency $f = 150 Hz$?


<!-- grade_id: cell-1e8745081b4f5d11, points: 0, grade: false, solution: true -->
```python
f = 150 # Hz
f_s = f*2  # Replace with correct value
print(f"{f_s = }")
```

Output:
```
f_s = 300
```

<!-- grade_id: cell-7de42bf2facd0b4e, points: 4, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(f_s, 0)).encode('utf-8')+b'97470').hexdigest() == '97941af8305597dd6055cef1f73fcc41b58d4327', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## e)
A analog signal $x_3(t) = \sin(2\pi \cdot f_3 \cdot t)$ with frequency $f_3 = 120 \text{ Hz}$ is sampled with a sampling frequency $f_s = 200 \text{ Hz}$ to create a digital signal $x_3[n]$. If we use the samples in $x_3[n]$ to attempt reconstruction of $x_3(t)$, what will the apparent frequency $\hat{f}_3$ be?

<!-- grade_id: cell-eb7a92c045bebdc6, points: 0, grade: false, solution: true -->
```python
f_3 = 120 # Hz
fs = 100 # Hz

# Calc the apparent frequency, while also checking for the correct slope of the signal.
f_3_est = (-(fs*(f_3//fs) - f_3)) * (not((f_3//fs) % 2)) + (fs - (f_3 - fs*(f_3//fs))) * ((f_3//fs) % 2)
    
print(f"{f_3_est = }")# Replace with correct value
```

Output:
```
f_3_est = 80
```

<!-- grade_id: cell-145b4d280d9b5fbb, points: 5, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(f_3_est), 0)).encode('utf-8')+b'e0b5a').hexdigest() == 'af4de889487697bb865df296f3d75c9de93e1c35', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## f) 
Given an ADC operating between $0.0\text{ V}$ and $3.3 \text{ V}$ and sampling with 12 bit resolution, what is the quantization step size $Q$ (aka. $\text{lsb value}$)?

*Use Volts as unit of measurement, even though the number is very small*

<!-- grade_id: cell-58aeaf630f3317a3, points: 0, grade: false, solution: true -->
```python
volt_range = 3.3
resolution = 2 ** 12

lsb_value = volt_range / resolution  # Replace with correct value
print(f"{lsb_value = }")
```

Output:
```
lsb_value = 0.0008056640625
```

<!-- grade_id: cell-c5a37d709348f434, points: 4, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(lsb_value), 6)).encode('utf-8')+b'652c3').hexdigest() == '6f603ebe37040d239cd9bb14e1c400afcf4d877e', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```

## g)
Given an ADC operating between $0.0\text{ V}$ and $3.3 \text{ V}$ and sampling with 12 bit resolution, and an input signal $x(t) = 1 + 0.7\cdot \cos(2\pi \cdot 50 \cdot t) \text{ V}$.

What is the Signal-to-Noise ratio (SNR) in decibels for the sampled signal? <br>
*PS. We will disregard the DC component, and only focus on the signal power for the AC component (there is no information to be extracted from a DC value).*

<!-- grade_id: cell-f6ab6e8109b08461, points: 0, grade: false, solution: true -->
```python
bits = 12
resolution = 2 ** bits
Vp = A = 0.7
volt_range = 3.3

q = lsb_value = volt_range / resolution # Quantization step / a.k.a the step length in volts
LF = 1/np.sqrt(2) # This is for full scale LF

signal = LF**2 * Vp ** 2
noise = (q**2)/12

SNR = 10*np.log10(signal/noise) # Replace with correct value

print(SNR)
```

Output:
```
66.5603934659193
```

<!-- grade_id: cell-7ceaf16dbc4d77d8, points: 5, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(float(SNR), 1)).encode('utf-8')+b'2b2b9').hexdigest() == '5a7d359a53821d630e73caff2e5ee908c974dd8a', 'Wrong answer :('
print('Correct answer :)')
```

Output:
```
Correct answer :)
```